<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Create a generator function `nmultiples` that returns a generator object that can be used to produce all of the multiples of a number N. Your function should be able to generate the 1,000,000th multiple in under a second.

Please note that the sequence in this problem starts with 0.

```python
>>> m = nmultiples(7)
>>> [next(m) for _ in range(5)]
[0, 7, 14, 21, 28]
```
[2 points]

<!-- grade_id: cell-32574a3fb8ed5982, points: 0, grade: false, solution: true -->
```python
def nmultiples(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2970ff4efe5245f7, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
m = nmultiples(7)
assert [next(m) for _ in range(5)] == [0, 7, 14, 21, 28]
### END TESTS
```

<!-- grade_id: cell-447405b27b670689, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-c4da759e55299384, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import inspect

assert inspect.isgeneratorfunction(nmultiples)
### END TESTS
```

**Q**: Create a generate that returns points along a noisy sine wave. The noise should not be added in the amplitude of the wave but in the values passed to `math.sin` . For example, the generator should produce `math.sin` of the `start_radians` given and then, for all following values, produce a value 0 to 0.2 radians (inclusive) ahead of the last value. To be clear, you can use the built in `math.sin` function to solve this problem and only add in noise via the value passed to that function. 

[2 points]

<!-- grade_id: cell-e8a9dfd93b0e0309, points: 0, grade: false, solution: true -->
```python
import math
import random

def noisy_signal(start_radians: float):
    radians = start_radians
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-59f5d59626ff94cd, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import inspect

assert inspect.isgeneratorfunction(noisy_signal)
### END TESTS
```

<!-- grade_id: cell-b0d0dd82643824ac, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert next(noisy_signal(4.5)) == -0.977530117665097
### END TESTS
```

<!-- grade_id: cell-c8b67e7d41fc06c4, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
gen = noisy_signal(0) 
points = [next(gen) for _ in range(100000)]
assert max(points) <= 1
assert min(points) >= -1
### END TESTS
```

<!-- grade_id: cell-6cb1eeb3900e89eb, points: 0.5, grade: true, solution: false -->
```python

```